# SLU00 - Jupyter Notebook: Exercise notebook

In this notebook you'll train the concepts learned in the Learning notebook. Good luck!

## Shortcuts
Run the cell below using two different shortcuts: first, use the shortcut that just rans the cell. Next, use the shortcut that runs the cell and inserts a cell below.   

```python
print("I'm a Jupyter Notebook master now")
```

Output:
```
I'm a Jupyter Notebook master now
```

```python
print("I'm a Jupyter Notebook master now")
```

Output:
```
I'm a Jupyter Notebook master now
```

```python

```

Use a shortcut to merge the 3 cells above this one:

```python
print("Please reunite me with the cell below")

print("I belong to everyone")

print("Please reunite me with the cell above")
```

## Notebook Kernel and memory 
The following cells are meant the illustrate the problem of Jupyter keeping the information in memory even when we delete the cells that were originally run with that information.
Run the 4 cells below. Delete the cell with the statement `b = 0.7` by using command mode and a shortcut. Run the 3 remaining cells again. Create a cell under the cell with the statement `print(a, b, c)` using the command mode and a shortcut. Print just the variable b. Now, run all cells again but restart the notebook first. Check that now you have an error when you try to print the variable b. After that, delete the print cell and go on to the next exercise: Markdown cells.

```python
a = 7
```

```python
b = 0.7
```

```python
c = 0.07
```

```python
print(a, b, c)
```

Output:
```
7 0.7 0.07
```

## Markdown cells
Create a markdown list in the markdown cell below, with 3 items: mathematics, physics and chemistry.

 - mathematics
 - physics
 - chemistry

Create, in the markdown cell below, a quote using markdown notation (tip: Google is your best friend)

> Hello World!

Change the cell below to markdown and create an header that says: "I've finished my first exercise notebook!"

# I've finished my first exercise notebook!

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and assign your slack id to a `slack_id` variable in the cell bellow.
Example:

```python
slack_id = "UTS63FC02"

<!-- grade_id: cell-db86bff08f67584f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
slackid = 'U038A3JJ0MQ'
```

<!-- grade_id: cell-8ea0aa7dc9e60a19, points: 0, grade: true, solution: false -->
```python
from submit import grade_submit

assert slackid is not None
#grade_submit(notebook_name='Exercise notebook.ipynb', learning_unit=0, exercise_notebook=1, slackid='migueldias')
```

```python
grade_submit(notebook_name='Exercise notebook.ipynb', learning_unit=0, exercise_notebook=1, slackid=slackid)
```

Output:
```
Success {"id":2839,"created":"2022-04-03T17:29:07.943606Z","slackid":"U038A3JJ0MQ","learning_unit":0,"exercise_notebook":1,"score":0.0}
```